### This is Assignment_0. The purpose of this assignment is to help you get familiar with assignment notebooks, datahub and assignment submission.

#### Submission Guidelines:
1. Save and run all the cells of this notebook before submitting.
2. Use the Validate button located on the jupyter notebook toolbar to check your solution. If the solution passes all the tests, a success pop-up will show. Otherwise, a pop-up will show which cells have failed.
4. You will be able to test your work on the local test cases by running the cell with assertion commands.
5. There are "Hidden Test Blocks" which will test the solutions on hidden test cases. Make sure to run all the blocks to ensure the test cases are tested.
6. Once you are done with your code, you should save the notebook and submit it in the assignment section on datahub. Ensure that you submit your assignment using the "Submit" button on the Assignment section.
7. To download the file as an .py, navigate to File -> Download as -> .py
8. Plagiarism and AI similarity check will be conducted.


##### Q1: Count Occurrences of Characters in a String
##### Complete the function 'char_count' that takes a string 's' and returns a dictionary with the count of each character in the string.


<!-- grade_id: cell-8b0dda2fe66089ee, points: 2, grade: true, solution: true -->
```python
def char_count(s):
    letter_count = {}
    for item in s:
        letter_count[item] = s.count(item)
    return(letter_count)
    raise NotImplementedError()
```

<!-- grade_id: cell-9f050990e62a15a3, points: 2, grade: true, solution: false -->
```python
assert char_count("data science") == {'d': 1, 'a': 2, 't': 1, ' ': 1, 's': 1, 'c': 2, 'i': 1, 'e': 2, 'n': 1}
```

<!-- grade_id: cell-b2fb9b09de8291f7, points: 2, grade: true, solution: false -->
```python
assert char_count("DSC 200") == {'D': 1, 'S': 1, 'C': 1, ' ': 1, '2': 1, '0': 2}
```

<!-- grade_id: cell-6f9c7668909cfde8, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-2e2b1983a2dba5e6, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2.
```

##### Q2: Filter Even Numbers from a List
##### Complete the function 'filter_even_numbers' that takes a list of integers 'nums' and returns a new list containing only the even numbers.


<!-- grade_id: cell-c4791a89eb1b344b, points: 2, grade: true, solution: true -->
```python
def filter_even_numbers(nums):
    even_num = []
    for item in nums:
        if item % 2 == 0:
            even_num.append(item)
        else:
            even_num
    return(even_num)
    raise NotImplementedError()
```

<!-- grade_id: cell-3e057532502b5f9b, points: 2, grade: true, solution: false -->
```python
assert filter_even_numbers([7, 8, 10, 13]) == [8, 10]
```

<!-- grade_id: cell-22bc7f10d0c30a33, points: 2, grade: true, solution: false -->
```python
assert filter_even_numbers([1, 3, 5]) == []
```

<!-- grade_id: cell-6d66833b2d347cfc, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-7246dbffee8dfedb, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2.
```

##### Q3: Find Common Elements Between Two Lists
##### Complete the function that takes two lists of integers ('list1' and 'list2') and returns a list of common elements between them.


<!-- grade_id: cell-0d0b22056d9df845, points: 2, grade: true, solution: true -->
```python
def common_elements(list1, list2):
    value = []
    for item1 in list1:
        for item2 in list2:
            if item1 == item2:
                value.append(item1)
            else: 
                value
    return(value)
    raise NotImplementedError()
```

<!-- grade_id: cell-878c32ea72004c55, points: 2, grade: true, solution: false -->
```python
assert set(common_elements([7, 8, 9], [10, 8, 7])) == {8, 7}
```

<!-- grade_id: cell-bcff1fa76baa7064, points: 2, grade: true, solution: false -->
```python
assert common_elements([1, 2, 3], [4, 5, 6]) == []
```

<!-- grade_id: cell-0a42fad61865ba6a, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-b195443e9af04608, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2.
```

```python

```

```python

```

```python

```