#### Submission Guidelines:
1. Save and run all the cells of this notebook before submitting.
2. Use the Validate button located on the jupyter notebook toolbar to check your solution. If the solution passes all the tests, a success pop-up will show. Otherwise, a pop-up will show which cells have failed.
4. You will be able to test your work on the local test cases by running the cell with assertion commands.
5. There are "Hidden Test Blocks" which will test the solutions on hidden test cases. Make sure to run all the blocks to ensure the test cases are tested.
6. Once you are done with your code, you should save the notebook and submit it in the assignment section on datahub. Ensure that you submit your assignment using the "Submit" button on the Assignment section.
7. To download the file as an .py, navigate to File -> Download as -> .py
8. Plagiarism and AI similarity check will be conducted.


### Assignment_1

##### Question 1:

##### Define a function called filter_large_cities that accepts a dictionary representing city names and their respective populations (in millions). The function should return a list of cities with populations greater than 1.5 million people, sorted in ascending order based on their population. Use a lambda function and the filter function to return the filtered list.

##### Example Input: {"San Diego": 1.38, "Los Angeles": 3.849, "New York": 8.468, "Austin": 0.964}
##### Expected Output: ["Los Angeles", "New York"]

<!-- grade_id: cell-8b0dda2fe66089ee, points: 2, grade: true, solution: true -->
```python
def filter_large_cities(d):
    """
    This function takes up a dictionary with key value pairs of cities with their population and returns
    a sorted list with cities more than 1.5 million population.
    """
    large_pop = lambda x: d[x]> 1.5 
    city_name = list(filter(large_pop, d))
    sorted_city_name = sorted(city_name, key = lambda x: d[x])
    return sorted_city_name
    raise NotImplementedError()
```

<!-- grade_id: cell-9f050990e62a15a3, points: 2, grade: true, solution: false -->
```python
# Test case 1
dict_cities = {"San Diego": 1.38, "Los Angeles": 3.849, "New York": 8.468, "Austin": 0.964}
assert filter_large_cities(dict_cities) == ["Los Angeles", "New York"]
```

<!-- grade_id: cell-b2fb9b09de8291f7, points: 2, grade: true, solution: false -->
```python
# Test case 2
dict_cities_2 = {"Chicago": 2.7, "Houston": 2.3, "Phoenix": 1.7, "Miami": 0.6}
assert filter_large_cities(dict_cities_2) == ["Phoenix", "Houston", "Chicago"]

```

<!-- grade_id: cell-6f9c7668909cfde8, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1. Don't alter the cell
```

<!-- grade_id: cell-2e2b1983a2dba5e6, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2. Don't alter the cell
```

##### Question 2:
##### Write a function called find_longest_unique_substring to find the length of the longest substring without repeating characters in a given string. For example, the longest substring without repeating characters in "banana" is "ban" or "ana" (if multiple same-length substrings are found, return the first one).
##### Example Input: "banana"
##### Expected Output: "ban"

<!-- grade_id: cell-c4791a89eb1b344b, points: 2, grade: true, solution: true -->
```python
def find_longest_unique_substring(s):
    """
    input: string
    output: longest substring without repeating characters
    """
    j = 0 
    sub_string = set()
    longest_sub_str = ""
    longest_length = 0
    for i in range(len(s)):
        while s[i] in sub_string:
            sub_string.remove(s[j])
            j+=1
        sub_string.add(s[i])
        current_length = i-j+1 
        if current_length > longest_length:
            longest_length = current_length 
            longest_sub_str = s[j:i+1]
    return(longest_sub_str)
    raise NotImplementedError()
```

<!-- grade_id: cell-3e057532502b5f9b, points: 2, grade: true, solution: false -->
```python
# Test case 1
assert find_longest_unique_substring("banana") == "ban"
```

<!-- grade_id: cell-22bc7f10d0c30a33, points: 2, grade: true, solution: false -->
```python
# Test case 2
assert find_longest_unique_substring("pwwkew") == "wke"
```

<!-- grade_id: cell-6d66833b2d347cfc, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1. Don't alter the cell
```

<!-- grade_id: cell-7246dbffee8dfedb, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2. Don't alter the cell
```

##### Question 3:
##### Write a function called solve_simple_maze that takes a 2D list representing a maze of size MxN. The maze starts at (0, 0) and ends at (M-1, N-1). You can move only to the right or down, through cells that are " " (empty), but cannot move through cells that are "X" (blocked).
##### There is only one path without dead ends, and you need to mark the path using the value "P" to indicate it is part of the solution. If the maze cannot be solved (if the start or end is blocked), return -1.

##### Example Input: [[" ", "X", " "], [" ", " ", " "], ["X", "X", " "]]
##### Expected Output: [["P", "X", " "], ["P", "P", "P"], ["X", "X", "P"]]


<!-- grade_id: cell-0d0b22056d9df845, points: 2, grade: true, solution: true -->
```python
def solve_simple_maze(maze):
    """
    input: 2D array
    output: solved 2D array
    """
    M = len(maze)
    N = len(maze[0])
    m = 0 #position in row
    n = 0 #position in column
    if maze[0][0] == "X" or maze[M-1][N-1] == "X":
        return -1
    while m < M and n < N:
        maze[m][n] = "P"
        move_right = False
        move_down = False
        if n+1 < N and maze[m][n+1] !="X":
            move_right = True
        if m+1 < M and maze[m+1][n] !="X":
            move_down = True
        if move_right == True:
            n +=1
        elif move_down == True:
            m +=1
        else:
            return -1
        if m == M-1 and n==N-1:
            maze[m][n] = "P"
            return maze
    return -1
    raise NotImplementedError()
```

<!-- grade_id: cell-878c32ea72004c55, points: 2, grade: true, solution: false -->
```python
# Test case 1
maze1 = [[" ", "X", " "], [" ", " ", " "], ["X", "X", " "]]
assert solve_simple_maze(maze1) == [["P", "X", " "], ["P", "P", "P"], ["X", "X", "P"]]

```

<!-- grade_id: cell-bcff1fa76baa7064, points: 2, grade: true, solution: false -->
```python
# Test case 2
maze2 = [[" ", "X"], ["X", " "]]
assert solve_simple_maze(maze2) == -1  # No possible path

```

<!-- grade_id: cell-0a42fad61865ba6a, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1. Don't alter the cell
```

<!-- grade_id: cell-b195443e9af04608, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2. Don't alter the cell
```

##### Question 4:
##### You are given a list of email addresses. Write a Python function called extract_clean_company_names that extracts and returns the unique company names from the email domains in the list.
 
##### Specifications:
##### The function should take a single argument emails, which is a list of email addresses (strings).
##### It should return a sorted list of unique company names extracted from the email domains in the input list.
##### The company name is defined as the part of the email domain between "@" and ".".
##### Remove any digits from the company name before returning.

##### **Example:**
##### emails = [
#####     "john.doe@example1.com",
#####     "jane_smith@company123.net",
#####     "alice@domain.com",
#####     "test.user@example.com",
#####     "user@domain2.net"
##### ]
 
##### **Returns**: ['company', 'domain', 'example']


<!-- grade_id: cell-3bdf9e9201b15cdf, points: 2, grade: true, solution: true -->
```python
import re 
def extract_clean_company_names(emails):
    """
    This function takes input as multiple email id's with different domain names and returns a sorted list of
    only unique domain names altering any digits in it.
    """
    company_list = set()

    for i in emails:
        company_name= re.findall('@([a-zA-Z]+)', i)
        company_list.add(company_name[0])
    return sorted(list(company_list))
    raise NotImplementedError()
```

<!-- grade_id: cell-1801ccb79144ed21, points: 2, grade: true, solution: false -->
```python
# Test case 1
emails_1 = [
    "john.doe@example1.com",
    "jane_smith@company123.net",
    "alice@domain.com",
    "test.user@example.com",
    "user@domain2.net"
]
result_1 = extract_clean_company_names(emails_1)
assert result_1 == ['company', 'domain', 'example'], "Test case 1 failed"

```

<!-- grade_id: cell-858916296f650fd0, points: 2, grade: true, solution: false -->
```python

# Test case 2 
emails_2 = [
    "john.doe@example99.com",
    "jane_smith@example.net",
    "alice@example3.com"
]
result_2 = extract_clean_company_names(emails_2)
assert result_2 == ['example'], "Test case 2 failed"

```

<!-- grade_id: cell-7ce24dc8a53c7543, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1. Don't alter the cell
```

<!-- grade_id: cell-8ca1dfe6891500f5, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2. Don't alter the cell
```

##### Question 5:
##### Write a function called generate_custom_sequence that generates a custom sequence using recursion of length n (provided as an argument).
 
 **Custom Sequence Pattern:**
 1. The sequence starts with the values **2** and **3**.
 2. From the **third element onward**, each element is calculated by multiplying the **last element** with the **second-to-last** element and then adding 2.
 3. The function must use **recursion**, and loops or other methods are **not allowed**.
 
 **Recursive Case:**
 - If n > 2, recursively generate the sequence of length n-1 and then calculate the next element using the last two elements.
 
 **Examples:**
 - generate_custom_sequence(4) should return [2, 3, 8, 26].
 - generate_custom_sequence(6) should return [2, 3, 8, 26, 210, 5462].
 
 **Explanation for n = 4:**
 1. **First two elements** are 2 and 3.
     - So, for n = 2: [2, 3]  
     
 2. **Third element:**
     - Use the formula: **Next Element = Last Element (3) * Second-to-Last Element (2) + 2**
     - Result: 3 * 2 + 2 = 8
     - So, for n = 3: [2, 3, 8]  
       
 3. **Fourth element:**
     - Again use the formula: **Next Element = Last Element (8) * Second-to-Last Element (3) + 2**
     - Result: 8 * 3 + 2 = 26
     - So, for n = 4: [2, 3, 8, 26]
 

<!-- grade_id: cell-2e593b21bda41eb8, points: 2, grade: true, solution: true -->
```python
def generate_custom_sequence(n):
    """
    input: integer n
    output: First n elements of the list based on the given pattern  
    """
    if n == 0:
        return False
    if n == 1:
        return [2]
    if n == 2: 
        return [2,3]
    else:
        previous_number_list = generate_custom_sequence(n-1)
        last_el = previous_number_list[-1]
        second_last_el = previous_number_list[-2]
        previous_number_list.append(last_el * second_last_el + 2)
        return previous_number_list
    raise NotImplementedError()
```

<!-- grade_id: cell-0738dd7481f3313b, points: 2, grade: true, solution: false -->
```python
# Test case 1
assert generate_custom_sequence(4) == [2, 3, 8, 26], "Test case 1 failed"

```

<!-- grade_id: cell-d9f6a5c1f695e0b5, points: 2, grade: true, solution: false -->
```python
# Test case 2
assert generate_custom_sequence(6) == [2, 3, 8, 26, 210, 5462], "Test case 2 failed"

```

<!-- grade_id: cell-3b9dd5dcd844f70f, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1. Don't alter the cell
```

<!-- grade_id: cell-8b31de529b09dbd5, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2. Don't alter the cell
```

```python

```

```python

```

```python

```